```python
# Initialize Otter
import otter
grader = otter.Notebook("ENGR131_Midterm.ipynb")
```

# ⌛️ Midterm

This midterm will assess your mastery of the course content in Weeks 1-4.


## Instructions

Note: For each question, your code should be indented under what is called the function definition: `def question_1():`. This helps the autograder run all the code you write by calling a single name. 

Note: Python cares about **Tabs**. Make sure spacing is correct in your functions. If you get an `Indention Error` it likely means your spacing is not correct. 

Your solution should look like this: 

```python
# DO NOT ALTER lines of code defining the function
def question ():
    line1 = line1_code # line 1 of code
    line2 = line2_code # line 2 of code

    # DO NOT ALTER lines of code in the return statement
    return x, y, z # variables to return separated by commas
```

**Your code replaces the prompt:  `...`**



## Part 1: Entering Your Information for Credit

To receive credit for assignments, it is important we can identify your work from others. To do this, we will ask you to enter your information in the following code block.

You must hit submit and see the message saying the data is saved. If you do not see this message contact your instructor to ensure you are getting credit for your work.

**Note: do not delete the output.log file or the .responses.json file or you might lose your current work.**

### Before you begin

Run the block of code at the top of the notebook that imports and sets up the autograder. This will allow you to check your work. 

```python
from ENGR131_Util_2024 import StudentInfoForm
from ENGR131_Util_2024 import responses
from ENGR131_Util_2024 import engr131_midterm

# Register the log function to be called before any cell is executed
get_ipython().events.register('pre_run_cell', engr131_midterm)

StudentInfoForm(**responses)
```

**Question 1: Identify python objects** 
*8 points*

Please run the block of code below and respond to the prompts with the most suitable response. 

Make sure you hit the submit button or your grade will not be recorded. Confirmation that your responses are recorded will appear below the code block. We will only grade your most recent set of responses. You can submit entries as many times as you like. 

```python
from ENGR131_Util_2024 import DataTypes
from ENGR131_Util_2024 import responses

self = DataTypes(**responses)

```

**Question 2: Create logical statements** 
*9 points*

Developing logical statements is essential to coding, especially controlling the flow of a program. 

Please run the block of code below and respond to the prompts with the most suitable response. 

Make sure you hit the submit button or your grade will not be recorded. Confirmation that your responses are recorded will appear below the code block. We will only grade your most recent set of responses. You can submit entries as many times as you like. 

```python
from ENGR131_Util_2024 import LogicalOperators
from ENGR131_Util_2024 import responses

logical_operators = LogicalOperators(**responses)
```

**Question 3: Unit conversion with a dictionary** 
*10 points*

As an engineer, you will often need to conduct unit conversions. For instance, a value measurement may be in feet that you need in inches or meters. 

A dictionary can be an efficient way to store conversion factors and access them for converting values from one unit to another. 

Write a function named `unit_conversion` that 
1. Does the conversion of the value and stores the result in a variable named `converted_value`.
1. Uses the dictionary `conversion_factors` to print the value to two decimal places with the original units and the converted value to two decimal places with the converted units. 

For instance, the function should print the following for the test input included below:

`10.00 meters is equal to 32.81 feet.`

**Do not print any other statements or the autograder will not work.**

```python
import numpy as np

# define a function called unit_conversion to be used for grading
def unit_conversion(value, from_unit, to_unit):
    # Define conversion factors for common units
    conversion_factors = {
        # Length
        'meters': {'meters': 1, 'feet': 3.28084, 'inches': 39.3701, 'centimeters': 100},
        'feet': {'meters': 0.3048, 'feet': 1, 'inches': 12, 'centimeters': 30.48},
        'inches': {'meters': 0.0254, 'feet': 0.0833333, 'inches': 1, 'centimeters': 2.54},
        'centimeters': {'meters': 0.01, 'feet': 0.0328084, 'inches': 0.393701, 'centimeters': 1},
        # Mass
        'kilograms': {'kilograms': 1, 'pounds': 2.20462, 'grams': 1000},
        'pounds': {'kilograms': 0.453592, 'pounds': 1, 'grams': 453.592},
        'grams': {'kilograms': 0.001, 'pounds': 0.00220462, 'grams': 1},
    }

    # Check if the units are in the conversion factors dictionary
    if from_unit in conversion_factors and to_unit in conversion_factors[from_unit]:
    
        ...
    else:
        print("Invalid units for conversion.")
        return None
        
    # The following line outputs the value from the function so that they can be accessed by the grader
    return converted_value

# Check the usage:
value_to_convert = 10
original_unit = 'meters'
target_unit = 'feet'

converted_value = unit_conversion(value_to_convert, original_unit, target_unit)
```

```python
grader.check("q3-unitConversion")
```

**Question 4: Calculate the vertical component of a force** 
*10 points*

Forces are often represented as vectors that can be broken down into horizontal and vertical components. The sine function is used when calculating the vertical component, especially when the force is applied at an angle.

For example, if we have a force `F` applied at an angle `θ` (in degrees) to the horizontal, the vertical component of the force `F_vertical` can be calculated as:

$$ F_{vertical} = F \cdot \sin(\theta) $$

Write python code to do the following inside the `question_4` function below:

* Define the magnitude of the force, `F`, as 100 Newtons (with a DataType INT) and the angle, `theta`, as 78 degrees (with a DataType FLOAT).
* Convert the angle `theta` to radians using the `math.radians` function because the `math.sin` function expects the angle in radians, save this to the variable `theta_radians`.
* Calculate the vertical component of the force, `F_vertical`, in Newtons.

```python
import numpy as np
import math

# define a function called question_4 to be used for grading
def question_4():

    ...

    # The following line outputs these values from the function so that they can be accessed by the grader
    return F, theta, theta_radians, F_vertical

F, theta, theta_radians, F_vertical = question_4()
print(sum([F,theta, theta_radians, F_vertical]))
```

```python
grader.check("q4-forceVector")
```

**Question 5: Using Python as a calculator** 

We have seen how Python can be used to evaluate mathematical expressions. This problem provides practice incorporating various mathematical functions and constants.  


Write python code to do the following:

* Define variables: `a = 3`, `b = -2`, and `x = 1`
* For each of the quantities $E_0$ through $E_2$, construct a one-line Python expression that computes the value and assigns it to a variable. (You should perform your computations using the variable names $a$, $b$, and $c$ instead of their numerical values). In Python, we will use variable names `E0` through `E2` to store the values $E_0$ through $E_2$. Definitions of the variables `E0` through `E2` have been started for you in the template.

$\text{i.  } E_0 = ln(4a+b)$

$\text{ii.  } E_1 = \cos \left( \left( \frac{\pi}{6a} \right)^2 \right)$

$\text{i.  } E_2 = \sqrt{x^2+a^2+b^2}$

You should use the `NumPy` library for the required mathematical functions. 

We have used the standard `NumPy` convention `import numpy as np`, thus, for example the `cos` function is used by typing `np.cos(<value>)`

Again, your code should be indented under the function definition: `def question_5():`.

Your code replaces the prompt:  `...`

```python
# import the np library to be used
import numpy as np

# define a function called question_3 to be used for grading
def question_5():

    a = ...
    b = ...
    x = ...

    E0 = ...
    E1 = ...
    E2 = ...

    # output these values from the function so that they can be accessed by the grader
    return a, b, x, E0, E1, E2
```

```python
grader.check("q5-calculator")
```

**Question 6: Analyzing Temperature Distribution Data** 

In thermal analysis, engineers often deal with temperature distribution across different materials or regions. For example, a mechanical engineer might analyze the temperature distribution across a metal plate, or a civil engineer might be interested in temperature variations across a building's surface.

You are given a randomly-populated, two-dimensional array representing temperature data, `temperature_grid`, with a length of 6 in each dimension.

Write Python code to do the following:
 * assign `specific_temperature` the temperature reading at the second row and fifth column.
 * assign `row_temperatures` the temperature readings in the fifth row.
 * assign `column_temperatures` the temperature readings in the second column.
 * assign `sub_matrix` the temperature readings in a 3 row by 2 column sub-matrix starting from the second row and third column.

Your code should replace this prompt: ...

```python
import numpy as np
 
def question_6(seed=None):
    # Set the seed for reproducibility, if provided
    if seed is not None:
        np.random.seed(seed)
 
    # Generate a 6x6 grid of temperature readings
    temperature_grid = np.random.rand(6, 6) 
 
    # Assign the temperature reading at the second row and fifth column
    specific_temperature = ...
 
    # Assign the temperature readings in the fifth row
    row_temperatures = ...
 
    # Assign the temperature readings in the second column
    column_temperatures = ...
 
    # Assign the temperature readings in a 3 row by 2 column sub-matrix
    # starting from the second row and third column
    sub_matrix = ...
 
    return specific_temperature, row_temperatures, column_temperatures, sub_matrix

```

```python
grader.check("q6-sensor-selection")
```

## Submitting Your Assignment

To submit your assignment please use the following link the assignment on GitHub classroom.
   
Use this [link](https://classroom.github.com/a/JkkvD3eK) to navigate to the assignment on GitHub classroom.

**Please submit just your output.log file. No other files are needed.**


If you need further instructions on submitting your assignment, please ask your TA during lab. 

## Viewing your score

Each `log` file you have uploaded will have a file with the name of your file + `Grade_Report.md`. You can view this file by clicking on the file name. This will show you the results of the autograder. 

**Make sure you run your code, and run the tests. If you do not run the test you will not get credit for your work.**

```{note}
In python and particularly jupyter notebooks it is common that during testing you run cells in a different order, or run cells and modify them. This can cause there to be local variables needed for your solution that would not be recreated on running your code again from scratch. Your assignment will be graded based on running your code from scratch. This means before you submit your assignment you should restart the kernel and run all cells. You can do this by clicking `Kernel` and selecting `Restart and Run All`. If you code does not run as expected after restarting the kernel and running all cells it means you have an error in your code. 
```

## Fin

Tests for q3-unitConversion (points: 10)
```
>>> import drexel_jupyter_logger
>>> value_to_convert = 12
>>> original_unit = 'meters'
>>> target_unit = 'feet'
>>> expected_output = drexel_jupyter_logger.decrypt_outputs('gAAAAABltVgJZoIAjULZM5ziXyLjuqhOWKak4onhEC1rOG_rTeZcjf1cD-cA0Q-g8qizRLRgdOAfxzRyxlrpTwXVeMinabTIR-lUvY0jUod1E44F2oH2F02qGhgXCvRvbi7vl64F_fe6')
>>> from unittest.mock import patch
>>> with patch('builtins.print') as mock_print:
...     conversion = unit_conversion(value_to_convert, original_unit, target_unit)
...     mock_print.assert_called_once_with(expected_output[:-1])
>>> if mock_print.call_args[0][0] == expected_output[:-1]:
...     drexel_jupyter_logger.variable_logger_csv('3, 3', 'q3_1')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 3', 'q3_1')
```
failure_message: Variables incorrectly defined.
success_message: Variables correctly defined.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> value_to_convert = 15
>>> original_unit = 'kilograms'
>>> target_unit = 'pounds'
>>> with patch('builtins.print') as mock_print:
...     conversion = unit_conversion(value_to_convert, original_unit, target_unit)
...     expected_output = drexel_jupyter_logger.decrypt_outputs('gAAAAABltVjOMtNIHXAQS0szhBAiLkYfz5rhyZuPO_7dbkP8uT2TVR9iD4fXn2-k1gjS8dYNSKzwML3abgNpbHa7lC6OciL_eBtTXYj8YLoXaA9hjc2dDjk3d0phENk2a5yy9ZD3Qcgf')
...     mock_print.assert_called_once_with(expected_output[:-1])
>>> if mock_print.call_args[0][0] == expected_output[:-1]:
...     drexel_jupyter_logger.variable_logger_csv('5, 5', 'q3_2')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 5', 'q3_2')
```
failure_message: Final answer is incorrect.
success_message: Final answer is correct.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> value_to_convert = 10
>>> original_unit = 'invalid_unit'
>>> target_unit = 'unknown_unit'
>>> condition_to_check = True
>>> expected_message = drexel_jupyter_logger.decrypt_outputs('gAAAAABltVo3gsLO9QaPGRMyLGs0uLtCX-9R1TGTVBxeQbKrzK8K4oxrH9RohAk7fcvpv4qFdHGSbCUv7tWT9VKrhZE_kpXn_tfviANrPlS0VZdeA1-SDIc=')
>>> with patch('builtins.print') as mock_print:
...     out = unit_conversion(value_to_convert, original_unit, target_unit)
...     mock_print.assert_called_once_with(expected_message[:-1])
>>> if mock_print.call_args[0][0] == expected_message[:-1]:
...     drexel_jupyter_logger.variable_logger_csv('2, 2', 'q3_3')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 2', 'q3_3')
```
failure_message: Final answer is incorrect.
success_message: Final answer is correct.


Tests for q4-forceVector (points: 14)
```
>>> import drexel_jupyter_logger
>>> import math
>>> (F, theta, theta_radians, F_vertical) = question_4()
>>> assert F == int(drexel_jupyter_logger.decrypt_outputs('gAAAAABltV0uIP98o5gsJM8AdcMCRMVGLCgbVB4QCFNf7c_nEg2BLY9xscg2PIxhsRGvR5ogXPrTyikwljacjYmd1i5xqARsTw==')) and theta == float(drexel_jupyter_logger.decrypt_outputs('gAAAAABltV3MyiRNnJOHmaS4ATEEDZwGeGWI-Gqjbz1227PZKqpAZ9g4OE8i5mlvg7fZPyxvoIy0r-4DES4QFWg6tgX8rqn82g=='))
>>> assert type(F) == int and type(theta) == float, 'DataTypes are not correct'
>>> exec(drexel_jupyter_logger.decrypt_outputs('gAAAAABluqYq6OBrYmm4KT7izVpEJbIRhdCrtyaM-mi6wE3jEj2MpksuPw0srLjuGdDVK8--pqMmKl-PYRK95DnSQMmMscnynPP7-ZaBDplRHu1g-VkdUr19OjD-FihBj794BPtkKBEta4Cpim4QxEvS9zHK3SqBcxKDdplxlydP3-BpbRbXC_YMBp50G6pj44J99bonCYtsWSPch8IjsqQOJ8nt6k0MeJBI4JnIxITrGcmLGtpTjqU-Qc5kZHaV62vG9Zpzlibv4PK6TGoXRenjpm_yNtEH3Q=='))
```
failure_message: Variables incorrectly defined.
success_message: Variables correctly defined.
```
>>> import numpy as np
>>> import drexel_jupyter_logger
>>> (F, theta, theta_radians, F_vertical) = question_4()
>>> value = float(drexel_jupyter_logger.decrypt_outputs('gAAAAABltV5yK5_Zr5VEAWOcL3nbPapn6Tf7l3wY-Sm3zIhvtydvp8sj5RbZwxisI5I6JEzA837dXywEWqQvPRj5Q3xiOnw2ocIwclYJ7bF9aTNAwLv6rz0='))
>>> assert np.isclose(sum([F, theta, theta_radians, F_vertical]), value)
>>> if np.isclose(sum([F, theta, theta_radians, F_vertical]), value):
...     drexel_jupyter_logger.variable_logger_csv('7, 7', 'q4_2')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 7', 'q4_2')
```
failure_message: Final answer is incorrect.
success_message: Final answer is correct.


Tests for q5-calculator (points: 15)
```
>>> (a, b, x, E0, E1, E2) = question_5()
>>> import drexel_jupyter_logger
>>> exec(drexel_jupyter_logger.decrypt_outputs('gAAAAABltV_GTkae8Fu2ykIHvm5AjnLGuz7WqJSGkv-cgHe7UiIAEpzPUOmeoR11X_xIoUCvoXnktW0R6KApzd4pJkbfCBd6zfMF_OATzv38_xJWwu4uIN1Ge2R64Ef85FGMlNzm9lz9'))
>>> exec(drexel_jupyter_logger.decrypt_outputs('gAAAAABluqaoGZJNEJJS4Y9M8FdWLk_9WyHuWS2sCeZ50J0Kot7EWQlw8ogR2_KG9-NUPeVMfZrwrZOnxkTzUpOT7tn4c4aoo73FKgaqMimeBRYhOLkXwTGz267D0IpacDyPFNbX8wQUu7OOVU8oMJ-HwjQIQyq5q0-cH_2l0Hg5ndye79SGONBtnmDYV-7J6MuQ66P5fl3_cGt9HRuoL_lBbBX1ZDeo-jmuWccfgNpvNoWIlon2ZURp-Mj7g9rVXBOrgBIenIBgpaspk56FudlajHCgzDWpJOtHyxiLP6PrJ1M4-6EYrPQ='))
```
failure_message: Variables incorrectly defined.
success_message: Variables correctly defined.
```
>>> import numpy as np
>>> import drexel_jupyter_logger
>>> (a, b, x, E0, E1, E2) = question_5()
>>> value = float(drexel_jupyter_logger.decrypt_outputs('gAAAAABltWG2z_X2afxl_mwinnL59AtnZk0gwoe-1QX4fagphlvcVkwl-PTH_kJZyABdXNOLrWSJkdmdareYKiYHED2EbKx2AEI1dY5EV8WZc-eJlPWSYHc='))
>>> assert np.isclose(sum([a, b, x, E0]), value)
>>> if np.isclose(sum([a, b, x, E0]), value):
...     drexel_jupyter_logger.variable_logger_csv('4, 4', 'q5_2')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 4', 'q5_2')
```
failure_message: E0 answer incorrect.
success_message: E0 answer correct.
```
>>> import drexel_jupyter_logger
>>> import numpy as np
>>> (a, b, x, E0, E1, E2) = question_5()
>>> value = float(drexel_jupyter_logger.decrypt_outputs('gAAAAABltWIsG_Z7exT1HiR31eMCs9VHYL24OIQCjcoZQsEimaOWgx1md1MEoUNr4oMY5AXtwVoGNhkKLH9r3l3_laduddxr_0ikjEEeT0LKVztBwRZznso='))
>>> assert np.isclose(sum([a, b, x, E1]), value)
>>> if np.isclose(sum([a, b, x, E1]), value):
...     drexel_jupyter_logger.variable_logger_csv('4, 4', 'q5_3')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 4', 'q5_3')
```
failure_message: E1 answer incorrect.
success_message: E1 answer correct.
```
>>> import drexel_jupyter_logger
>>> import numpy as np
>>> (a, b, x, E0, E1, E2) = question_5()
>>> value = float(drexel_jupyter_logger.decrypt_outputs('gAAAAABltWLJdY__9xAY3fAwQStd08TJT3RQSHV84zEq1VO6jdT-EJPZVrqKBOwLEdpPmNuZC1yhh_yZz_lohrKhMFZ_-07cVn_5QgVIic1--ica0agi01o='))
>>> assert np.isclose(sum([a, b, x, E2]), value)
>>> if np.isclose(sum([a, b, x, E2]), value):
...     drexel_jupyter_logger.variable_logger_csv('4, 4', 'q5_4')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 4', 'q5_4')
```
failure_message: E2 answer incorrect.
success_message: E2 answer correct.


Tests for q6-sensor-selection (points: 16)
```
>>> import drexel_jupyter_logger
>>> (specific_temperature, row_temperatures, column_temperatures, sub_matrix) = question_6(42)
>>> assert type(specific_temperature) == np.float64
>>> assert len(row_temperatures) == 6
>>> assert len(column_temperatures) == 6
>>> assert sub_matrix.shape == (3, 2)
>>> if type(specific_temperature) == np.float64 and len(row_temperatures) == 6 and (len(column_temperatures) == 6) and (sub_matrix.shape == (3, 2)):
...     drexel_jupyter_logger.variable_logger_csv('4, 4', 'q6_1')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 4', 'q6_1')
```
success_message: Function correctly returns types.
```
>>> import drexel_jupyter_logger
>>> (specific_temperature, row_temperatures, column_temperatures, sub_matrix) = question_6(42)
>>> value = np.float32(drexel_jupyter_logger.decrypt_outputs('gAAAAABluqczpmrtih9-tXyof9yWCMDoFQSIrXI7hG5UfzfEOkPLNK8lA71Vz0gI0LW1cOos7o5MVDtvxKCmL0j7yQiz3Uixyt1GpV_yPYJokSs_YObnzzg='))
>>> assert np.isclose(specific_temperature, value)
>>> if np.isclose(specific_temperature, value):
...     drexel_jupyter_logger.variable_logger_csv('3, 3', 'q6_2')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 3', 'q6_2')
```
success_message: specific_sensor_reading is correct.
```
>>> import drexel_jupyter_logger
>>> (specific_temperature, row_temperatures, column_temperatures, sub_matrix) = question_6(42)
>>> value = np.float32(drexel_jupyter_logger.decrypt_outputs('gAAAAABluqeN0r0kLNRA4dIgCdJDHhgdBCFrZi5xgaqCzLyljmlwqbLmZwOesRPU5WJ0Xp3jMC2bj73OqSfRu8cuGaSQwiwF29S-VREEDBvMgK_wWUzj5k8='))
>>> assert np.isclose(np.sum(row_temperatures), 2.5940191477640613)
>>> if np.isclose(np.sum(row_temperatures), 2.5940191477640613):
...     drexel_jupyter_logger.variable_logger_csv('3, 3', 'q6_3')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 3', 'q6_3')
```
success_message: row is correct.
```
>>> import drexel_jupyter_logger
>>> (specific_temperature, row_temperatures, column_temperatures, sub_matrix) = question_6(42)
>>> value = np.float32(drexel_jupyter_logger.decrypt_outputs('gAAAAABluqe01srohfN8lNwhVJPKhLz6rA9ZR4Xs6BPMvo38JQl4GkZML7cGILxDoQnlRZoEfEPscS3OW3P-k8eHodIsImUNLmJz-pko_eJKPfCoGo1d8nU='))
>>> assert np.isclose(np.sum(column_temperatures), value)
>>> if np.isclose(np.sum(column_temperatures), value):
...     drexel_jupyter_logger.variable_logger_csv('3, 3', 'q6_4')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 3', 'q6_4')
```
success_message: row is correct.
```
>>> import drexel_jupyter_logger
>>> (specific_temperature, row_temperatures, column_temperatures, sub_matrix) = question_6(42)
>>> value = np.float32(drexel_jupyter_logger.decrypt_outputs('gAAAAABluqfmFxQiu701pbPjjBW5Hyifly7yIhuyZlY_P3HsqrwG4W1ISpJJ-oO5qUZrYZ9MACe1qGI6I8HyJY232jrUjlLYOXl5eKdGMVH_5SLHz6PpT4U='))
>>> assert np.isclose(np.sum(sub_matrix), value)
>>> if np.isclose(np.sum(sub_matrix), value):
...     drexel_jupyter_logger.variable_logger_csv('3, 3', 'q6_5')
... else:
...     drexel_jupyter_logger.variable_logger_csv('0, 3', 'q6_5')
```
success_message: row is correct.
